## Imports

```python
# pip install --upgrade scikit-learn
```

```python
import time
from IPython.display import clear_output
import numpy    as np
import pandas   as pd
import seaborn  as sb
import matplotlib.pyplot as plt
import sklearn  as skl

from sklearn.pipeline import Pipeline      
from sklearn import pipeline
from sklearn.preprocessing import OrdinalEncoder, LabelEncoder
from sklearn.impute import SimpleImputer
from sklearn import compose
from sklearn.model_selection import train_test_split
from sklearn import model_selection
from sklearn import metrics
from sklearn.metrics import accuracy_score, balanced_accuracy_score, plot_confusion_matrix         # accuracy_score, balanced_accuracy_score, plot_confusion_matrix
from sklearn import set_config
from sklearn.compose import ColumnTransformer


set_config(display='diagram') # Useful for display the pipeline

print("Pandas  ", pd.__version__)
print("Sklearn ", skl.__version__) # Try to use 0.24
```

Output:
```
Pandas   1.4.1
Sklearn  1.0.2
```

## Get the dataset
- **CLOUD = True**: Download dataset from Kaggle. Necesary for cloud enviroments like COLAB. **Specify your [kaggle credentials](https://www.kaggle.com/docs/api)**.
- **CLOUD = False**: Get the dataset from your local machine. **Specify the data path**.

```python
DATA_PATH = "data/"
```

## Load data

```python
print(DATA_PATH)
```

Output:
```
..08. Robust ML/data
```

```python
df      = pd.read_csv(DATA_PATH + "train.csv", index_col='PassengerId')
df_test = pd.read_csv(DATA_PATH + "test.csv",  index_col='PassengerId')

print("Train DataFrame:", df.shape)
print("Test DataFrame: ", df_test.shape)
```

Output:
```
Train DataFrame: (891, 11)
Test DataFrame:  (418, 10)
```

## Check missings

```python
df.isnull().sum()
```

Output:
```
Survived      0
Pclass        0
Name          0
Sex           0
Age         177
SibSp         0
Parch         0
Ticket        0
Fare          0
Cabin       687
Embarked      2
dtype: int64
```

```python
df_test.isnull().sum()
```

Output:
```
Pclass        0
Name          0
Sex           0
Age          86
SibSp         0
Parch         0
Ticket        0
Fare          1
Cabin       327
Embarked      0
dtype: int64
```

# Exercise 1 (2pts):
Extract the title (Mr, Mrs, ... ) from the "Name" column.

Tips:
- split(',')[1] to get the 2nd part, and remove the surnamename
- split('.')[0] to get the 1str part, and remove the name

<!-- grade_id: cell-d56107842f0f54d2, points: 0, grade: false, solution: true -->
```python
# CODE HERE get_Title_from_Name funtion
# Create this function using lambda (not def)
import re
# def name_title(df):
#     df['Title'] = [i.split(',')[1].split('.')[0] for i in df['Name'].to_list()]
#     return df['Title']
    
get_Title_from_Name = lambda x: x.split(',')[1].split('.')[0].strip()
df['Title'] = df['Name'].apply(get_Title_from_Name)
df_test['Title'] = df_test['Name'].apply(get_Title_from_Name)


df['Title'].values[886]
```

Output:
```
'Rev'
```

<!-- grade_id: cell-f3d989a17018344b, points: 2, grade: true, solution: false -->
```python
assert df['Title'].values[0] == "Mr"
assert df['Title'].values[1] == "Mrs"
assert df['Title'].values[2] == "Miss"

assert df_test['Title'].values[0] == "Mr"
assert df_test['Title'].values[1] == "Mrs"
assert df_test['Title'].values[414] == "Dona"
```

# Exercise 2 (1pts):
Apply the title_dictionary to get a better information about the title. You have to overwrite the Title variable.

```python
title_dictionary = {
    "Capt": "Officer",
    "Col": "Officer",
    "Major": "Officer",
    "Jonkheer": "Royalty",
    "Don": "Royalty",
    "Sir" : "Royalty",
    "Dr": "Officer",
    "Rev": "Officer",
    "the Countess":"Royalty",
    "Mme": "Mrs",
    "Mlle": "Miss",
    "Ms": "Mrs",
    "Mr" : "Mr",
    "Mrs" : "Mrs",
    "Miss" : "Miss",
    "Master" : "Master",
    "Lady" : "Royalty"
}
```

<!-- grade_id: cell-5c6a842c812fafda, points: 0, grade: false, solution: true -->
```python
# Use map to apply the prevous dict
# titles = df['Title']
# titles_test = df_test['Title']
df["Title"] = df["Title"].map(title_dictionary) 
df_test["Title"] = df_test["Title"].map(title_dictionary) 



```

<!-- grade_id: cell-669f5a637e57e835, points: 1, grade: true, solution: false -->
```python
assert df['Title'].values[886] == "Officer"
assert df_test['Title'].values[417] == "Master"
```

# Exercise OPTINAL (0pts):
Try to extract some information from the feature **Ticket**. Search on Internet if that colum has some kind of information.

```python

```

Output:
```
PassengerId
892                 330911
893                 363272
894                 240276
895                 315154
896                3101298
               ...        
1305             A.5. 3236
1306              PC 17758
1307    SOTON/O.Q. 3101262
1308                359309
1309                  2668
Name: Ticket, Length: 418, dtype: object
```

# Exercise OPTIONAL (0pts):
Try to extract some information from the feature **Cabin**. Search on Internet if that colum has some kind of information.

```python

```

Output:
```
PassengerId
892      NaN
893      NaN
894      NaN
895      NaN
896      NaN
        ... 
1305     NaN
1306    C105
1307     NaN
1308     NaN
1309     NaN
Name: Cabin, Length: 418, dtype: object
```

# Preprocessing
For X data, notice that...
- We drop Survived because is the target variable
- We drop Name because we have extracted the Title: Mr, Mrs, ...
- We drop Ticket because it has no information -> see df.Ticket.nunique()
- We drop Cabin because it has a lot of missings (77% are missings)

Then, we identify **numerical** variables and **categorical** variables,

```python
x = df.drop(columns=["Survived", 'Name', 'Ticket', 'Cabin', 'Titles']) # X DATA (WILL BE TRAIN+VALID DATA)
y = df["Survived"] # 0 = No, 1 = Yes

x_test = df_test.drop(columns=['Name', 'Ticket', 'Cabin']) # # X_TEST DATA (NEW DATA)
```

```python
cat_vars  = ['Sex', 'Embarked', 'Title']         # x.select_dtypes(include=[object]).columns.values.tolist()
num_vars  = ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age'] # x.select_dtypes(exclude=[object]).columns.values.tolist()

print("\nNumerical features:\n", num_vars)
print("\nCategorical features:\n", cat_vars)
```

Output:
```
Numerical features:
 ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age']

Categorical features:
 ['Sex', 'Embarked', 'Title']
```

# Exercise 3 (2pts):
Create a **ColumnTransformer for Tree Models**. You need to create 2 pipelines (one for numerical and other for categories). Remember:
- Categorical pipeline: Some SimpleImputer -> Some Encoder
- Numerical pipeline: Some SimpleImputer -> NO Encoder

```python
x
```

Output:
```
Pclass     Sex   Age  SibSp  Parch     Fare Embarked    Title
PassengerId                                                               
1                 3    male  22.0      1      0   7.2500        S       Mr
2                 1  female  38.0      1      0  71.2833        C      Mrs
3                 3  female  26.0      0      0   7.9250        S     Miss
4                 1  female  35.0      1      0  53.1000        S      Mrs
5                 3    male  35.0      0      0   8.0500        S       Mr
...             ...     ...   ...    ...    ...      ...      ...      ...
887               2    male  27.0      0      0  13.0000        S  Officer
888               1  female  19.0      0      0  30.0000        S     Miss
889               3  female   NaN      1      2  23.4500        S     Miss
890               1    male  26.0      0      0  30.0000        C       Mr
891               3    male  32.0      0      0   7.7500        Q       Mr

[891 rows x 8 columns]
```

<!-- grade_id: cell-c607e75fb38ea248, points: 0, grade: false, solution: true -->
```python
"""
num_preprocessing = pipeline.Pipeline(steps=[
  # Some SimpleImputer here
])

cat_preporcessing = pipeline.Pipeline(steps=[
  # Some SimpleImputer here
  # Some Encoder here. Remember to handle_unknown
])

tree_prepro = compose.ColumnTransformer(transformers=[
    ('num', num_preprocessing, num_vars),
    ('cat', cat_preporcessing, cat_vars),
], remainder='drop') # Drop other vars not specified in num_vars or cat_vars

tree_prepro
""";
### BEGIN SOLUTION
num_4_treeModels = Pipeline( [ ('imputer', SimpleImputer()) ] )
cat_4_treeModels = Pipeline( [ ('imputer', SimpleImputer(strategy='most_frequent')), ('ordinal', OrdinalEncoder(handle_unknown= 'use_encoded_value', unknown_value = -1))])
tree_prepro = ColumnTransformer( [ ('num', num_4_treeModels, num_vars), ('cat', cat_4_treeModels, cat_vars) ], remainder='drop' )
tree_prepro

```

Output:
```
ColumnTransformer(transformers=[('num',
                                 Pipeline(steps=[('imputer', SimpleImputer())]),
                                 ['Pclass', 'SibSp', 'Parch', 'Fare', 'Age']),
                                ('cat',
                                 Pipeline(steps=[('imputer',
                                                  SimpleImputer(strategy='most_frequent')),
                                                 ('ordinal',
                                                  OrdinalEncoder(handle_unknown='use_encoded_value',
                                                                 unknown_value=-1))]),
                                 ['Sex', 'Embarked', 'Title'])])
```

<!-- grade_id: cell-e636558135fb5975, points: 2, grade: true, solution: false -->
```python
assert type(tree_prepro)      is compose._column_transformer.ColumnTransformer
assert type(num_4_treeModels) is pipeline.Pipeline
assert type(cat_4_treeModels) is pipeline.Pipeline
assert len(num_4_treeModels) == 1
assert len(cat_4_treeModels) == 2
```

# Exercise 4 (1pts):
1. Complete the diccionary with some Tree Models.
2. Then we put each model in a Pipeline where:
   - first is the prepocessing with the column Transformer
   - Then is the Tree model
3. Display the fullpipeline of the LGBMClassifier

```python
from sklearn.tree          import DecisionTreeClassifier
from sklearn.ensemble      import RandomForestClassifier
from sklearn.ensemble      import ExtraTreesClassifier
from sklearn.ensemble      import AdaBoostClassifier
from sklearn.ensemble      import GradientBoostingClassifier
from sklearn.ensemble      import HistGradientBoostingClassifier
from xgboost               import XGBClassifier
from lightgbm              import LGBMClassifier
from catboost              import CatBoostClassifier
```

<!-- grade_id: cell-76a3be8223730c5a, points: 0, grade: false, solution: true -->
```python
"""
tree_classifiers = {
  "Decision Tree": DecisionTreeClassifier(),
  "Extra Trees":
  "Random Forest":
  "AdaBoost":
  "Skl GBM":
  "Skl HistGBM":
  "XGBoost":
  "LightGBM":
  "CatBoost":
tree_classifiers = {name: pipeline.make_pipeline(tree_prepro, model) for name, model in tree_classifiers.items()}
tree_classifiers["LightGBM"]
"""
tree_classifiers = {
  "Decision Tree": DecisionTreeClassifier(random_state=0),
  "Extra Trees": ExtraTreesClassifier(random_state=0),
  "Random Forest": RandomForestClassifier(random_state=0),
  "AdaBoost": AdaBoostClassifier(random_state=0),
  "Skl GBM": GradientBoostingClassifier(random_state=0),
  "Skl HistGBM": HistGradientBoostingClassifier(random_state=0),
  "XGBoost": XGBClassifier(),
  "LightGBM": LGBMClassifier(random_state=0),
  "Catboost": CatBoostClassifier(random_state=0)
}
tree_classifiers = {name: pipeline.make_pipeline(tree_prepro, model) for name, model in tree_classifiers.items()}
tree_classifiers["LightGBM"]
```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('num',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer())]),
                                                  ['Pclass', 'SibSp', 'Parch',
                                                   'Fare', 'Age']),
                                                 ('cat',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(strategy='most_frequent')),
                                                                  ('ordinal',
                                                                   OrdinalEncoder(handle_unknown='use_encoded_value',
                                                                                  unknown_value=-1))]),
                                                 
```

<!-- grade_id: cell-d4744022b37e2e9b, points: 1, grade: true, solution: false -->
```python
for pipe in tree_classifiers.values():
    assert type(pipe) is pipeline.Pipeline
```

# Exercise 5 (3pts):
Define a simple split validation strategy with:
- 80% for train
- 20% for validation
- With stratification
- random_state=0

And train all the models in a for loop

<!-- grade_id: cell-64f17e0d448bca7e, points: 0, grade: false, solution: true -->
```python
x_train, x_val, y_train, y_val = train_test_split(x, y, test_size=0.2, stratify=y, random_state=0)


results = pd.DataFrame({'Model': [], 'Accuracy': [], 'Bal Acc.': [], 'Time': []})


for model_name, model in tree_classifiers.items():
    start_time = time.time()
    
    # FOR EVERY PIPELINE (PREPRO + MODEL) -> TRAIN WITH TRAIN DATA (x_train)
    model.fit(x_train, y_train)

    
    # GET PREDICTIONS USING x_val
    pred = model.predict(x_val)
    

    total_time = time.time() - start_time

    results = results.append({"Model":    model_name,
                              "Accuracy": metrics.accuracy_score(y_val, pred)*100,
                              "Bal Acc.": metrics.balanced_accuracy_score(y_val, pred)*100,
                              "Time":     total_time},
                              ignore_index=True)
                              
                              


### BEGIN SOLUTION




results_ord = results.sort_values(by=['Accuracy'], ascending=False, ignore_index=True)
results_ord.index += 1 
# results_ord.style.bar(subset=['Accuracy', 'Bal Acc.'], vmin=0, vmax=100, color='#5fba7d')
```

Output:
```
C:\Users\Lenovo\AppData\Local\Temp\ipykernel_3576\3046484079.py:20: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  results = results.append({"Model":    model_name,
C:\Users\Lenovo\AppData\Local\Temp\ipykernel_3576\3046484079.py:20: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  results = results.append({"Model":    model_name,
C:\Users\Lenovo\AppData\Local\Temp\ipykernel_3576\3046484079.py:20: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  results = results.append({"Model":    model_name,
C:\Users\Lenovo\AppData\Local\Temp\ipykernel_3576\3046484079.py:20: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  results = results.append({"Model":    
```

```python
results
```

Output:
```
Model   Accuracy   Bal Acc.      Time
0  Decision Tree  77.094972  75.421607  0.032952
1    Extra Trees  77.653631  75.065876  0.270266
2  Random Forest  78.212291  75.520422  0.267659
3       AdaBoost  79.888268  77.964427  0.177522
4        Skl GBM  84.916201  82.595520  0.150596
5    Skl HistGBM  81.005587  78.603426  1.675674
6        XGBoost  81.564246  78.787879  0.259813
7       LightGBM  81.005587  78.333333  0.334780
8       Catboost  81.564246  79.057971  2.810907
```

<!-- grade_id: cell-af27b33fe88ad384, points: 3, grade: true, solution: false -->
```python
assert results_ord["Accuracy"].min() > 75
assert results_ord["Bal Acc."].min() > 75
assert len(results_ord) == 9
```

# Exercise 6 (3pts):
Define a 10 Fold cross validation strategy with:
- With stratification
- shuffle=True
- random_state=0

And train all the models in a for loop.

Tip you can use **[cross_val_predict](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.cross_val_predict.html)** for both training and predict with 

```python
"""
skf = model_selection.StratifiedKFold(
    # CODE HERE
)
"""
skf = model_selection.StratifiedKFold(shuffle=True, random_state=0)


results = pd.DataFrame({'Model': [], 'Accuracy': [], 'Bal Acc.': [], 'Time': []})


for model_name, model in tree_classifiers.items():
    start_time = time.time()
        
    # TRAIN AND GET PREDICTIONS USING cross_val_predict() and x,y
    pred = model_selection.cross_val_predict(model, x, y, cv=skf)

    total_time = time.time() - start_time

    results = results.append({"Model":    model_name,
                              "Accuracy": metrics.accuracy_score(y, pred)*100,
                              "Bal Acc.": metrics.balanced_accuracy_score(y, pred)*100,
                              "Time":     total_time},
                              ignore_index=True)
                              
                              



results_ord = results.sort_values(by=['Accuracy'], ascending=False, ignore_index=True)
results_ord.index += 1 
# results_ord.style.bar(subset=['Accuracy', 'Bal Acc.'], vmin=0, vmax=100, color='#5fba7d')
```

Output:
```
C:\Users\Lenovo\AppData\Local\Temp\ipykernel_3576\778780458.py:20: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  results = results.append({"Model":    model_name,
C:\Users\Lenovo\AppData\Local\Temp\ipykernel_3576\778780458.py:20: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  results = results.append({"Model":    model_name,
C:\Users\Lenovo\AppData\Local\Temp\ipykernel_3576\778780458.py:20: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  results = results.append({"Model":    model_name,
C:\Users\Lenovo\AppData\Local\Temp\ipykernel_3576\778780458.py:20: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  results = results.append({"Model":    mode
```

```python
print(results)
```

Output:
```
Model   Accuracy   Bal Acc.       Time
0  Decision Tree  77.104377  75.632729   0.229588
1    Extra Trees  80.471380  79.081584   1.030400
2  Random Forest  80.583614  79.062410   1.177806
3       AdaBoost  80.695847  79.539354   0.725318
4        Skl GBM  82.716049  80.682581   0.722739
5    Skl HistGBM  82.491582  80.555556   5.794395
6        XGBoost  81.481481  80.066628   1.559970
7       LightGBM  81.818182  80.174480   1.123992
8       Catboost  82.603816  80.701754  15.203286
```

<!-- grade_id: cell-57cc085faad513cf, points: 3, grade: true, solution: false -->
```python
assert results_ord["Accuracy"].min() > 75
assert results_ord["Bal Acc."].min() > 75
assert len(results_ord) == 9
```

# Exercise 7.1
Train with all data the best model

```python
# best_model = tree_classifiers["SELECT MY BEST MODEL HERE"]

# Fit best model with all data
best_model = tree_classifiers.get("Skl GBM")

best_model.fit(x_train, y_train)

```

Output:
```
Pipeline(steps=[('columntransformer',
                 ColumnTransformer(transformers=[('num',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer())]),
                                                  ['Pclass', 'SibSp', 'Parch',
                                                   'Fare', 'Age']),
                                                 ('cat',
                                                  Pipeline(steps=[('imputer',
                                                                   SimpleImputer(strategy='most_frequent')),
                                                                  ('ordinal',
                                                                   OrdinalEncoder(handle_unknown='use_encoded_value',
                                                                                  unknown_value=-1))]),
                                                 
```

# Exercise 7.2 (2pts)
With your best model, generate the predicitions for test data (x_test)

<!-- grade_id: cell-44b93a0dbd4eb6fc, points: 0, grade: false, solution: true -->
```python
test_pred = best_model.predict(x_test)


```

<!-- grade_id: cell-a29fc8d24284520e, points: 2, grade: true, solution: false -->
```python
assert len(test_pred) == 418
assert np.unique(test_pred).tolist() == [0,1]
```

# Exercise 7.3

Submit to kaggle.

- You can use the kaggle command line app. Check https://github.com/Kaggle/kaggle-api

```python
"""sub = pd.DataFrame(test_pred, index=x_test.index, columns=["Survived"])
sub.head()"""
```

Output:
```
Survived
PassengerId          
892                 0
893                 0
894                 0
895                 0
896                 1
```

```python
'''sub.to_csv("sub.csv")'''
```

```python
!kaggle competitions submit -c titanic -f sub.csv -m "My submission message"
```

Output:
```
'kaggle' is not recognized as an internal or external command,
operable program or batch file.
```